Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joshua Zempoalteca"
COLLABORATORS = "N/A"
```

---

# **Problem Set 4** Execution Control 
### This is a problem set. You can work on it with the help of your classmates, Google, or ask me any questions you would like.
### I encourage discussion during these problem sets. Todays problem set is relatively short.
### The problem set is due at 2:50 pm. 

```python
import numpy as np
from numpy import random 
```

### 1.The factorial of an integer is the product of the all the integers starting from the integer going down to one.  So, for example, in math we usually write 
### $ n!  = n * (n-1) * (n-2)....2 * 1 $
### for example, 5! = 5 * 4 * 3 * 2 * 1 = 120 
### of course in any real program, we would use np.math.factorial(n) to compute the factorial of an integer n. 
### But, we are here to suffer and learn.  
### Write a `while` loop to compute the factorial of any integer **n**, and save that output in a variable called **fact** 

<!-- grade_id: cell-2db16ee01ff4f885, points: 0, grade: false, solution: true -->
```python
n = 5 
fact = 1

while n > 1:
    fact *= n
    n -= 1
print(fact)
```

Output:
```
120
```

<!-- grade_id: cell-5e2489665ee39ae7, points: 3, grade: true, solution: false -->
```python
assert fact == np.math.factorial(5) 
```

### 2. Given the lengths of 3 sides of a triangle - x, y and z - use a statement of the type (`if`-`elif`-`else`) to determine whether the triangle is 'equilateral', 'isosceles' or 'obtuse' and set the variable **triangle** to the correct string value
### Note: equilateral means all sides are equal, isosceles means two of the sides are equal but not the third one, obtuse means all 3 sides are different.
### There is no preferred way to do this problem, and there are many possible solutions. 
### Note that i have created 3 test cases here. The code should be the identical for all 3 cases and produce the correct answer, and you should just copy and paste the same code.  
### When testing a piece of code, you should always test that all possible cases work. 

<!-- grade_id: cell-fbe274a9b7384feb, points: 0, grade: false, solution: true -->
```python
x = 6 
y = 7
z = 7

if (x == y):
    if x == z:
        triangle = 'equilateral'
    else:
        triangle = 'isosceles'
elif (x == z):
    if x == y:
        triangle = 'equilateral'
    else:
        triangle = 'isosceles'
elif (y == z):
    if y == x:
        triangle = 'equilateral'
    else:
        triangle = 'isosceles'
else:
    triangle = 'obtuse'
```

<!-- grade_id: cell-a014d87e7e27b765, points: 1, grade: true, solution: false -->
```python
assert triangle == 'isosceles'
```

<!-- grade_id: cell-052f0d28106ad644, points: 0, grade: false, solution: true -->
```python
x = 7 
y = 7
z = 7

if (x == y):
    if x == z:
        triangle = 'equilateral'
    else:
        triangle = 'isosceles'
elif (x == z):
    if x == y:
        triangle = 'equilateral'
    else:
        triangle = 'isosceles'
elif (y == z):
    if y == x:
        triangle = 'equilateral'
    else:
        triangle = 'isosceles'
else:
    triangle = 'obtuse'
```

<!-- grade_id: cell-39b6b44e41742e02, points: 1, grade: true, solution: false -->
```python
assert triangle == 'equilateral'
```

<!-- grade_id: cell-89074cd7c92bfe45, points: 0, grade: false, solution: true -->
```python
x = 6 
y = 7
z = 8

if (x == y):
    if x == z:
        triangle = 'equilateral'
    else:
        triangle = 'isosceles'
elif (x == z):
    if x == y:
        triangle = 'equilateral'
    else:
        triangle = 'isosceles'
elif (y == z):
    if y == x:
        triangle = 'equilateral'
    else:
        triangle = 'isosceles'
else:
    triangle = 'obtuse'
```

<!-- grade_id: cell-fa97d8e69f6e66bc, points: 1, grade: true, solution: false -->
```python
assert triangle == 'obtuse'
```

### 3.  Write a program that finds all the factors of a number **n** (positive integer) and saves them to a list called **factors**. If the number is prime, set a variable **isprime** to True, otherwise **isprime** is false
### An integer **m** is a factor of **n** if **n**/**m** is an integer (no remainder).  Hint: the % operator computes the remainder when dividing integers.
### There are two test cases to run.  Use the identical code for both test cases.  

<!-- grade_id: cell-60545d34375ada7a, points: 0, grade: false, solution: true -->
```python
#Test case 1
n = 99 

factors = []
isprime = True

for i in range(1, n+1):
    m = n%i
    if m == 0:
        factors.append(i)
        if i != 1 and i != n:
            isprime = False
```

<!-- grade_id: cell-232887bca53c71e8, points: 2, grade: true, solution: false -->
```python
assert isprime == False
assert len(factors) == 6
```

<!-- grade_id: cell-71e094f10279b9d0, points: 0, grade: false, solution: true -->
```python
#Test Case 2
n = 83

factors = []
isprime = True

for i in range(1, n+1):
    m = n%i
    if m == 0:
        factors.append(i)
        if i != 1 and i != n:
            isprime = False
```

<!-- grade_id: cell-1763a18efe2f06e3, points: 2, grade: true, solution: false -->
```python
assert isprime == True
assert len(factors) == 2
```