Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joshua Zempoalteca"
COLLABORATORS = ""
```

---

# **Problem Set 3** While and For Loops
### This is a problem set.  You can work on it with the help of your classmates, Google, or ask me any questions you would like. 
### I encourage discussion during these problem sets.  Todays problem set is relatively short.  
### The problem set is due at 2:40 pm.  We will spend the last 10 minutes of class briefly looking over my answers. 

```python
import numpy as np
from numpy import random 
from matplotlib import pyplot as plt
```

### 1.  Write a for loop that will print out the name of each animal in the list.  The number of items in the list should be called **nanimals** and obtained by using the length `len` function.  

<!-- grade_id: cell-8c32f649996e6487, points: 0, grade: false, solution: true -->
```python
animals =["koala", "cat", "fox", "panda", "chipmunk", "sloth", "penguin", "dolphin"]
naminals = 0
for i in range(len(animals)):
    print(animals[i])
    naminals += 1

```

Output:
```
koala
cat
fox
panda
chipmunk
sloth
penguin
dolphin
```

<!-- grade_id: cell-fe6280b80828dea5, points: 2, grade: true, solution: false -->
```python
assert len(animals) == 8
```

### 2. Create a list of **n** random integers (between 0 and 9), called **numbers**, using the given random number generator  **rng** 
### write a `for` loop to compute a *numpy array* **numbersum**  the cumulative sum of these numbers.  
* ### The first element of **numbersum** is the sum of the first element of **number**.  
* ### The second element of **numbersum** is the sum of the first 2 elements of **number**
* ### The third element of **numbersum** is the sum of the first 3 elements of **number** 
### .... (keep going)
### Your example should be run with **n** = 20 

<!-- grade_id: cell-3ab7bdfb0565b2a5, points: 0, grade: false, solution: true -->
```python
rng = random.default_rng(seed = 1111)
n = 20
numbers = rng.integers(0,10,n)

numbersum = np.zeros(20)

for i in range(len(numbers)):
    numbersum[i] = numbersum[i-1] + numbers[i]

```

<!-- grade_id: cell-0b00e9427adcf969, points: 4, grade: true, solution: false -->
```python
assert n == 20
assert np.size(numbers) == n
assert np.all(numbersum == np.cumsum(numbers)) 
assert numbersum.dtype == 'float64'
```

## 3. GEOMETRIC SERIES 

### $$ geom = \sum_{n=1}^{\infty}{\frac{1}{2^n}} = \frac{1}{2} + \frac{1}{4} + \frac{1}{8} + \frac{1}{16} + .....$$

### is a geometric series, that sums to 1 as **n** approaches infinity ($\infty$). 

### Write a `while` loop that will keep adding terms to the sum **geom** (i.e., increasing **n**) until the increase in the value of the sum  is smaller than a user-specified value (**tol**).  
### For your test code, use tol = 0.0001

### The code should return the number of terms **n** in the sum, and a variable called **error** which is the difference between **geom** and 1.   
### Hint: there should be a variable tested against **tol** by the while loop and updated on each iteration of the loop. 

<!-- grade_id: cell-a076f6a85f9e3f7f, points: 0, grade: false, solution: true -->
```python
tol = 0.0001

geom = 0
n = 0
error = 1

# Compute the sum until the error is smaller than the tolerance
while abs(error) > tol:
    n += 1
    term = 1/(2**n)
    geom += term
    error = geom - 1
    
```

<!-- grade_id: cell-889b5403193aac70, points: 4, grade: true, solution: false -->
```python
assert np.abs(error) < 1E-04
assert n == 14
```

```python

```