Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joshua Zempoalteca"
COLLABORATORS = "NA"
```

---

# Homework 1  (5 pts)

### This homework is a mini homework with a few examples of how to work with python variables, expression, and lists, in order to test the jupyter lab environment in each of your computers, and my automatic grading systems. 
### In order for automatic grading and validation to work, it is very important that you make use of the **exact variable names** specified in the problem statement.  
### I will always place these variable names that must be used for automatic grading in bold 
### Let's get started! 

### 1. Create a variable named **My_First_Variable** and set its value to 2.  (1 pt)

<!-- grade_id: cell-2745506731a213ef, points: 0, grade: false, solution: true -->
```python
My_First_Variable = 2
```

### We test solutions using assert statements.  For fun, try to do this problem correctly and incorrectly and see what it does when you push the validation button above.  

<!-- grade_id: cell-a1e9fdf4667bd8d6, points: 1, grade: true, solution: false -->
```python
assert My_First_Variable == 2
```

### 2. Create a variable **v** and set its value equal to 0.5. Compute the variable **w** using the following equation $\textbf{w} = 2\frac{1-\textbf{v}}{1+\textbf{v}}$ (1 pt)
### A useful technique here is to make use of parenthesis to control the order of operations. If you place (1-v) and and (1+v) insider parenthesis, they will be operated on first, before dividing.  
### You can also feel free to make intermediate variables e.g., x = 1-v and y = 1+v if that helps you organize the solution.  

<!-- grade_id: cell-989cfd2c6b242c38, points: 0, grade: false, solution: true -->
```python
v = 0.5
w = 2*((1-v)/(1+v))
```

<!-- grade_id: cell-2d6405744fe5ae36, points: 1, grade: true, solution: false -->
```python
assert v == 0.5
assert w == 2/3
```

### You can run the validate at any point.  Try it now.  

### 3. Create a list named **days** containing the days of the week, starting on 'Monday' and ending on 'Sunday' (1 pts)

<!-- grade_id: cell-f5f3ef2fa3517087, points: 0, grade: false, solution: true -->
```python
days = [ "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday", "Sunday" ]
```

<!-- grade_id: cell-e97c023c69ed93a6, points: 1, grade: true, solution: false -->
```python
assert len(days) == 7
assert days[0] == 'Monday'
assert days[6] == 'Sunday'
```

### 4. Use indexing to create two new lists, **Weekend** which contain 'Saturday' and 'Sunday', and **Weekdays** which contain the other 5 working days of the week. (2 pts)

<!-- grade_id: cell-cd3015af9040c09b, points: 0, grade: false, solution: true -->
```python
Weekend = [ "Saturday", "Sunday" ]
Weekdays = ["Monday", "Tuesday", "Wednesday", "Thursday", "Friday"]
```

<!-- grade_id: cell-1794b59379065020, points: 1, grade: true, solution: false -->
```python
assert len(Weekend) == 2
```

<!-- grade_id: cell-507689ec563fd49b, points: 1, grade: true, solution: false -->
```python
assert len(Weekdays) == 5
```