Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joshua Zempoalteca"
COLLABORATORS = "N/A"
```

---

# Homework 4

## Please answer the questions below by entering and executing code in the box below each question.  

## The goals of this assignment are to reinforce lessons about loops and execution control.  


```python
import numpy as np
```

### 1. Grading Problem.  
### I suck at using Canvas or even Excel, and really at anything where I cant write some code to solve my problems.   
### So, I am going to have to write some code to compute your final grades for the class, and convert it to a letter grade. 
### Your grade is based on 4 categories - 
* #### Classwork 20% 
* #### Tests 30%
* #### Homework 30%
* #### Final Project 20%
### As an example here are the grades for 5 students in each cateory
* classwork = [80, 95, 85, 55, 95]
* tests = [75, 90,90, 70,95]
* homework = [100, 90,95,90,95]
* project = [90,85,100,70,100]
### (1) Write code to compute the final grade for each student using the percentages given above in a list or array called **finalgrade**. 
### (2) Convert the final grade to a letter grade for each student in a list called **lettergrade**. Use the following rule to convert the grades - 
* finalgrade >= 90   A 
* finalgrade >= 80   B
* finalgrade >= 70   C
* finalgrade < 70    F 


```python
classwork = [80, 95, 85, 55, 95]
tests = [75, 90,90, 70,95]
homework = [100, 90,95,90,95]
project = [90,85,100,70,100]
```

<!-- grade_id: cell-1878b9d5af44d167, points: 0, grade: false, solution: true -->
```python
finalgrade = []

for student in range(len(classwork)):
    finalgrade.append((classwork[student] * .20) + (tests[student] * .30) + (homework[student] * .30) + (project[student] * .20))

lettergrade = []
    
for grade in range(len(finalgrade)):
    if finalgrade[grade] >= 70 and finalgrade[grade] < 80:
        lettergrade.append('C')
    elif finalgrade[grade] >= 80 and finalgrade[grade] < 90:
        lettergrade.append('B')
    elif finalgrade[grade] >= 90:
        lettergrade.append('A')
    else:
        lettergrade.append('F')
```

<!-- grade_id: cell-072d960d9255fdab, points: 1, grade: true, solution: false -->
```python
assert finalgrade[0] == 86.5 
assert np.max(finalgrade) == 96
```

<!-- grade_id: cell-0ed36a7abaaf1378, points: 2, grade: true, solution: false -->
```python
assert lettergrade[1] == 'A'
```

### 2. COLLATZ CONJECTURE

### This is one of these weird things with numbers that remains unproven, but always works. Collatz had this insight in 1939, but could not prove it, hence it is a conjecture. 
### Start with any positive integer **n**.   **n** should be the first element of your list. 
### Follow this rule: 
### (1) if the integer is even, the next integer is one half of the current integer. 
### (2) if the integer is odd, the next integer is 3 times the integer plus one. 
### Eventually the sequence of integers will reach 1. This always works! 
### Write some python code that generates such a sequence, and terminates at 1. Some hints. 
### (1) A while loop should test if the current state of the integer, and exit the loop once the sequence of integers has reached 1. 
### (2) Remember that the conditional statement n%2 == 0 will return **True** if n is even and **False** if n is odd.
### You code should save the sequence of integers in a list called **coll** 
### For the test case, we will start with **n** = 7


<!-- grade_id: cell-ed0c8fa36ef11a33, points: 0, grade: false, solution: true -->
```python
n = 7

coll = [n] 

while n != 1:
    if n%2 == 0:
        n = n // 2  
    else:
        n = 3 * n + 1
    coll.append(n)  
```

<!-- grade_id: cell-7853de86aa0f8d23, points: 3, grade: true, solution: false -->
```python
assert len(coll) == 17
assert coll[0] == 7
assert coll[-1] == 1
```

### 3. Approximating $\pi$.  There are a number of approximations to the number $\pi$ that have been developed over the years.   
### Leibniz's formula for approximating $\pi$ is

### $$ \pi = 4 - 4/3 + 4/5 - 4/7 + 4/9 - 4/11 + 4/13 ....   $$

### $$ \pi = \sum_{n=0}^{\infty}{(-1)^n\frac{4}{2n+1}} $$
### Note two things to pay attention to about this sum. 
* #### The first term is equivalent to 4/1.  So, we are counting 1,3,5,7,9,11.... for the denominator.  That is where the 2*n+1 comes from 
* #### We are alternating between adding and subtracting numbers.  That is where the $(-1)^n$ comes from.    

### a. Write a for loop that will approximate $\pi$ with N = 20 terms.  Compute a variable called **error** which is the difference between the exact value of $\pi$ and your approximation.  To get the exact value of pi from python use `np.pi` 


<!-- grade_id: cell-4ca8811b9f090f7e, points: 0, grade: false, solution: true -->
```python
N = 10
pi_approx = 0
for n in range(N):
    pi_approx += (-1)**n * 4 / (2*n+1)

error = np.pi - pi_approx
print("Approximation of pi:", pi_approx)
print("Error:", error)

```

Output:
```
Approximation of pi: 3.0418396189294032
Error: 0.09975303466038987
```

<!-- grade_id: cell-f04759fe654a63bc, points: 2, grade: true, solution: false -->
```python
assert np.abs(error) < 0.1 
```

### b.  Write a 'while' loop that sums the terms in Leibniz's formula for $\pi$ until the *change* in the sum is less than a user specified value **tol**.  For your test, use **tol** = 0.01.  After you complete the 'while' loop compute the value of **error** - the difference between your approximation and the exact value of pi 

<!-- grade_id: cell-ff3bb4f4a1331216, points: 0, grade: false, solution: true -->
```python
tol = 0.01

pi_approx = 0
n = 0
change = 1

while change >= tol:
    term = (-1)**n * 4 / (2*n+1)
    pi_approx += term
    change = abs(term)
    n += 1

error = np.pi - pi_approx
```

<!-- grade_id: cell-e43486eb0e3d3cc1, points: 2, grade: true, solution: false -->
```python
assert np.abs(error) < 0.005
```

```python

```