Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joshua Zempoalteca"
COLLABORATORS = "NA"
```

---

# Homework 5 

## Practice some of the fundamentals of writing and using functions


```python
import numpy as np
```

### 1.  Write a function called **circle_area** that computes the area of a circle given its radius **r**, using this formula
### $$ A = \pi r^2 $$
### The function should have one input argument (**r**) and return the area **A**
### The input argument **r** should take a default value of 1.  
### After writing and executing the definition of the function, test that it works with the three tests below 

<!-- grade_id: cell-bc214dcc4799e196, points: 0, grade: false, solution: true -->
```python
def circle_area(r):
    A = (np.pi * (r**2))
    return A

```

<!-- grade_id: cell-81fb41c51dd1ce6c, points: 1, grade: true, solution: false -->
```python
r = 2
Area  = circle_area(r) 
assert Area == np.pi*r**2
```

<!-- grade_id: cell-cb2ff476b366949d, points: 1, grade: true, solution: false -->
```python
Area = circle_area()
assert Area == np.pi*1**2 
```

<!-- grade_id: cell-6c767c3dc3d35aed, points: 1, grade: true, solution: false -->
```python
Area = circle_area(r=3)
assert Area == np.pi*3**2
```

### 2.  Write a function called **cylinder_area** that computes the surface area of a the cylinder given input arguments, radius (**r**) and height (**h**).  
### The surface area of the cylinder is the area of the top and the bottom (both circles of radius **r**) added to area of the side of the cylinder 
### The area of the side of the cylinder is given as $2 \pi rh$ where **h** is the height of the cylinder 
### Your new function **cylinder_area** should make use of the function **circle_area** to get the area of the top and bottom of the cylinder.  

<!-- grade_id: cell-2a065046e1f3fb63, points: 0, grade: false, solution: true -->
```python
def cylinder_area(r,h):
    top_bottom_area = 2 * circle_area(r)
    side_area = 2 * np.pi * r * h
    surface_area = top_bottom_area + side_area
    return surface_area

```

<!-- grade_id: cell-610acb29e99efa65, points: 1, grade: true, solution: false -->
```python
assert cylinder_area(r=2,h=1) == 12*np.pi
```

<!-- grade_id: cell-17668daeb1dca0f0, points: 1, grade: true, solution: false -->
```python
assert cylinder_area(r=1,h=2) == 6*np.pi
```

## 3. Distance function in arbitrary dimensions.  
### In a plane, a point a has location defined by a pair of values $(a_1,a_2)$ which gives its position along the (x,y) axes
###  The distance between two points in the plane a and b is then 
### $$ d = \sqrt{(b_2-a_2)^2 + (b_1-a_1)^2} $$
### If the point a is in 3 dimensions 3 coordinates are required, $(a_1,a_2,a_3)$ corresponding to (x,y,z).
### In 3 dimensions, the distance formula between two points a and b is 
### $$ d = \sqrt{(b_3-a_3)^2 + (b_2-a_2)^2 + (b_1-a_1)^2)} $$
### In this case each of a,b is an array of size 3.
### The notion of "distance" can be applied in any number of dimensions.  So, in general distance between 2 points in a an n-dimensional space can be defined as 
### $$ d = \sqrt{\sum_{j=1}^{n}{(b_j-a_j)^2}} $$ 
### Write a function called distance that takes as input arguments the two points a,b and returns the distance d between the points.  Each point is given as an array or list of size n.  
### The same function should be able to work for any array size n = 1,2,3,4......  

<!-- grade_id: cell-4497a6342694c829, points: 0, grade: false, solution: true -->
```python
def distance(a,b):
    a = np.array(a)  # Convert to numpy array
    b = np.array(b)  # Convert to numpy array
    
    diff = b - a  
    squared_diff = np.square(diff)  
    sum_squared_diff = np.sum(squared_diff)  
    d = np.sqrt(sum_squared_diff)
    return d

```

<!-- grade_id: cell-ecf84493493de182, points: 1, grade: true, solution: false -->
```python
assert distance(1,2) == 1
```

<!-- grade_id: cell-632c73c5bd0d5785, points: 1, grade: true, solution: false -->
```python
assert distance([2,3],[-1,-1]) == 5
```

<!-- grade_id: cell-2f97802ace3b0b8a, points: 1, grade: true, solution: false -->
```python
assert distance(np.ones(4), np.zeros(4)) == 2
```

### 4. The `input` function in python allows for keyboard input. 
### For example: 
###     d = input('test input:')
### will prompt for an input with the text 'test input'  and place the keyboard input in a variable d of type string. (Even if you enter a number it will create a variable of type string). 
### Write a function **get_response** which prompts a user for a response of either yes or no using an `input` statement with the following prompt:
### "Respond with either yes or no:"
### The function should check if the keyboard input was 'yes' or 'no'.  If the keyboard input is either 'yes' or 'no' it should return the response, otherwise it should print the statement 
### "Invalid Response, Try Again" 
### and prompt for a response again.  


<!-- grade_id: cell-7553c472a5fcaa23, points: 0, grade: false, solution: true -->
```python
def get_response():
    response = input("Respond with either yes or no: ")
    if response.lower() == 'yes':
        return response
    elif response.lower() == 'no':
        return response
    else:
        print("Invalid Response, Try Again")

```

<!-- grade_id: cell-e6d2c4f64560afbf, points: 2, grade: true, solution: false -->
```python
response = get_response()
assert (response == 'yes') | (response == 'no')
```

```python

```