Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joshua Zempoalteca"
COLLABORATORS = "python23*"
```

---

# **Test 2** 
### Please answer each question by writing the code that solves the question.  
### You may make use of any class notes you want.
### You may not use any internet searches during the test. 
### You may not use your phone during the test. 
### You may not communicate with anybody during the test.  

```python
import numpy as np
from numpy import random 
from matplotlib import pyplot as plt
```

### 1.  (3 pts)  write a for loop that prints a countdown starting from the number **n** to the number 1. For your test case, set n = 10     

<!-- grade_id: cell-059da929b42afd65, points: 0, grade: false, solution: true -->
```python
n = 10

while n > 1:
    print(n- 1)
    n -= 1
```

Output:
```
9
8
7
6
5
4
3
2
1
```

<!-- grade_id: cell-91b48bfcb562d1be, points: 3, grade: true, solution: false -->
```python
assert 1 == 1
```

### 2. (4 pts)  In your first year on a job you earn 50,000.  Every year your salary goes up by 4%.  Make a for loop that saves a list **salary** that has your salary for every year starting at 50000.  For your test case calculate the salary upto year 10 (inclusive).   

<!-- grade_id: cell-50afb01193ca4b3c, points: 0, grade: false, solution: true -->
```python
n = 10
loop_counter = 0
salary_total = 50000
salary = [salary_total]

while loop_counter < n:
    salary_total = salary_total + (salary_total * 0.04)
    salary.append(salary_total)
    loop_counter += 1
    
print(salary)
```

Output:
```
[50000, 52000.0, 54080.0, 56243.2, 58492.928, 60832.64512, 63265.9509248, 65796.588961792, 68428.45252026369, 71165.59062107424, 74012.2142459172]
```

<!-- grade_id: cell-66047205a9176432, points: 4, grade: true, solution: false -->
```python
assert len(salary) == 11 
assert salary[0] == 50000
assert salary[2] == 54080
```

### 3. (4 pts) The converse of the Pythagorean Theorem is:

### If the square of the length of *any* one side of a triangle is equal to the sum of the squares of the other two sides, then the triangle is a right triangle.
### Write an if-else or if-elif-else statement that uses the lengths of the three sides of the triangle ΔABC, with lengths **a**, **b**, **c** 
### and tests if a triangle is a right triangle using the Pythagorean theroem. If the triangle is a right triangle, set the variable **triangle** to the value 'right', or else set triangle to the value 'oblique'. 
### For your test case, use a=5, b=4, c=3. 

<!-- grade_id: cell-77a146f7e1ec6100, points: 0, grade: false, solution: true -->
```python
a = 5
b = 4
c = 3 

if a**2 == ((b**2) + (c**2)):
    triangle = 'right'
elif b**2 == ((c**2)+ (a**2)):
    triangle = 'right'
elif c**2 == ((a**2) + (c**2)):
    triangle = 'right'
else:
    triangle =  'oblique'
```

<!-- grade_id: cell-f6954365e5a536ec, points: 4, grade: true, solution: false -->
```python
assert triangle == 'right'
```

### 4. (4 pts) Write a while loop that counts the number of times you can divide a positive integer **n** by an integer **m**, without any remainder. 
### For example, the number 120 can be divided by the number 2, three times 

* 120/2 = 120/2**1 = 60
* 120/2/2 = 120/2**2 = 30 
* 120/2/2/2 = 120/2**3 = 15
* 120/2/2/2/2 = 120/2**4 = 7.5 (or 7 with a remainder of 1)

### The variable that keeps track of the number of times you can divide by 2 should be called **count**. In the case 
### Run your code on a test case where **n** = 405 and **m** = 3 

<!-- grade_id: cell-523c6d798aff5d23, points: 0, grade: false, solution: true -->
```python
n = 405
m = 3
count = 0

while n % m == 0:
   n = n // m
   count += 1
    
```

<!-- grade_id: cell-c4dcdc3ec577488e, points: 4, grade: true, solution: false -->
```python
assert count == 5
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[23], line 1
----> 1 assert count == 5

AssertionError:
```