Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Joshua Zempoalteca"
COLLABORATORS = ""
```

---

# **Problem Set 1**:  List, Arrays, and Indexing Methods

### This is a problem set.  You can work on it with the help of your classmates, Google, or ask me any questions you would like. 
### There are some new methods introduced in the problem set. 
### I encourage discussion during these problem sets.  Todays problem set is relatively short.  
### The problem set is due at 2:40 pm.  We will spend the last 10 minutes of class briefly looking over my answers.  

```python
# your going to need numpy to do this assignment 
import numpy as np 
#frankly, in this class, you will always need numpy.  
```

1. ### Create a variable **v** which is a list that contains the values 7, 16, 15, 14, 22.  Create a variable **w** which contains the values -1,-0.5,0,0.5,1 

<!-- grade_id: cell-73c2a41fa91bd822, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
v = [7, 16, 15, 14, 22]
w = [-1,-0.5,0,0.5,1]
```

<!-- grade_id: cell-8d1778e1145a2f3c, points: 1, grade: true, solution: false -->
```python
assert len(v) == 5
assert len(w) == 5
assert v[2] == 15
assert w[2] == 0 
```

2. ### Create a new variable **x** which is the product of 14 and -0.5, by appropriately indexing into **v** and **w**.   

<!-- grade_id: cell-18b63f3fdc24973e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x = v[-2] * w[1]
```

<!-- grade_id: cell-221caebb0c8f6642, points: 1, grade: true, solution: false -->
```python
assert x == -7
```

3. ### Create a new list **y** which contains the first 3 elements of v and a new list **z** which containes the middle 3 entries (2nd to 4th) of w.  


<!-- grade_id: cell-263425207e7a0916, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
y = v[0:3]
z = w[1:4]
```

<!-- grade_id: cell-3b992638a13bd153, points: 1, grade: true, solution: false -->
```python
assert len(y) == 3
assert len(z) == 3
```

4. ### Merge **y** and **z** to make a new list **yz** which contains the entries of **y** followed by the entries of **z**


<!-- grade_id: cell-d5b94f102e7b6267, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
yz = y + z
```

<!-- grade_id: cell-3f99586982b85245, points: 1, grade: true, solution: false -->
```python
assert yz[3] == w[2]
```

## **INDEXING TO THE END OF A LIST**  
### When indexing into a list or array, sometimes we want to get all the elements from a certain starting point to the end of the array. 
### OF course this is easy if you know how many elements there are in the array. 
### Python supports indexing to the end of an array. 
### e.g., to start from the 3rd element and go to the end of **v**, 

```python
v3 = v[2:] #notice i provide the starting index (2) but no ending index. 
print(v3)
```

Output:
```
[15, 14, 22]
```

## **INDEXING BACKWARDS FROM THE END OF A LIST**  
### When indexing into list or arrays, we sometimes want to get the last element, or a certain number of elements of the list or array.  
### Of course if you know how long the array is in advance, you can easily solve this.
### Python supports indexing from the end of an array using negative indexes. 
### e.g., to get the last element of w 

```python
last_w = w[-1]
print(last_w)

```

Output:
```
1
```

### To get the last 3 elements of w 

```python
last3_w = w[-3:]
print(last3_w)
```

Output:
```
[0, 0.5, 1]
```

5. ### Using negative indexing, make a new variable **last_vw** which is a new list constructed by merging the last *2* elements of **v** with the last *4* elements of **w**

<!-- grade_id: cell-6af9da420e7531f2, points: 0, grade: false, solution: true -->
```python
last_vw = v[-2:] + w[-4:]
print(last_vw)
```

Output:
```
[14, 22, -0.5, 0, 0.5, 1]
```

<!-- grade_id: cell-5363f5900c1f82b9, points: 1, grade: true, solution: false -->
```python
assert last_vw[0] == 14
assert last_vw[-1] == 1
```

6. ### Convert the lists **v** and **w** into a numpy array with the same variable name using the `array` function.  Compute a new array **vxw** by multiplying **v** and **w**.  

<!-- grade_id: cell-e09606133e1769d4, points: 0, grade: false, solution: true -->
```python
v = np.array(v)
w = np.array(w)
vxw = v * w
```

<!-- grade_id: cell-0e5bf107194cd0d1, points: 1, grade: true, solution: false -->
```python
assert vxw[0] == -7
assert vxw[-1] == 22 
```

7. Create a numpy array named **z** with 10 integers, all with value -100 

<!-- grade_id: cell-740ba0b12d446160, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
z = np.linspace(-100,-100, 10)
```

<!-- grade_id: cell-205673e4c080b846, points: 1, grade: true, solution: false -->
```python
assert z.size == 10
assert z[0] == -100
```

8. Create an array named **digit_2** containing the whole numbers from 10 to 99, including 99

<!-- grade_id: cell-85940fb721af412d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
digit_2 = np.arange(10,100,1)
```

<!-- grade_id: cell-f9e6ed8f985e86d6, points: 1, grade: true, solution: false -->
```python
assert digit_2[0] == 10
assert digit_2[-1] == 99
```

9. Make a new array **small_2**  which copies all of the values in **digit_2** between 10 and 20 using indexing, including 10 but *not* including 20.  
Make another array called **big_2** which copies all of the values in **digit_2** between 90 and 99 using indexing including 90 and 99.  

<!-- grade_id: cell-5f68b1d27ee7f588, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
small_2= digit_2[0:10]
big_2 = digit_2[80:90]
```

<!-- grade_id: cell-75d2b48eda06f112, points: 1, grade: true, solution: false -->
```python
assert small_2[0] == 10
assert small_2[-1] == 19
assert big_2[0] == 90
assert big_2[-1] == 99
```

10. Create an array named **line** using linspace with 15 linearly spaced points between 0 and 1. 

<!-- grade_id: cell-16acdffd385a34b5, points: 0, grade: false, solution: true -->
```python
line = np.linspace(0,1,15)
```

<!-- grade_id: cell-500152db6b1295d6, points: 1, grade: true, solution: false -->
```python
assert line[0] == 0
assert line[-1] == 1
assert line.size == 15
```